```python
# Initialize Otter
import otter
grader = otter.Notebook("PracticeAssignment.ipynb")
```

# Assignment 0: Practice - Gradescope

This assignment exists for practice only, and consist 0% of your grade.


This assignment covers assignment completion, _gradescope_ submission, and anaconda downloading.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**


## How to complete assignments

Whenever you see:

```python
answers = ...
```

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `...`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `grader.check("Question_id")` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. Once you run the cell, you will see the public tests we provided for you.

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

For each assignment, you will encounter following types of question -- multiple choice, math, coding, and free response. For each type of question, we expect you to answer them according to the format we asked so that the autograder will function properly.

### Q1 - Multiple Choice Question (0.1 points)

Consider your response to the following question: which course are you MOST excited about this quarter?:
- A) COGS 118B
- B) Introduction to Supervised Machine Learning
- C) COGS1111118A
- D) Intro Sup ML

Replace `...` below to assign the letter corresponding to your response to the object `favorite_class`. Be sure that you put quotes around the the letter for your response. 

For example, if "COGS 118A" (option A above) were my answer, I would replace `favorite_class = ...` with `favorite_class = 'A'`.

After you add your answer, run the code in that cell using 'Shift + Enter' or clicking the "Run" icon above. A number will appear to the left of the cell once your code has executed.

```python
# any answer is a good answer here!
favorite_class = 'A'
```

Typically, after you write some code in your assignments, there will be a few assert statements, similar to what you see below. These help to check that you're on the right path. If you are on the right track, when, you run the code in those cells, they will "pass silently", meaning you won't see any output. 

However, if you get an AssertionError, you'll want to return to the code you right and re-think your response.

_Points:_ 0.1

```python
grader.check("q1_multiple_choice_question")
```

Output:
```
q1_multiple_choice_question results: All test cases passed!
```

**Note**: Just because an assert doesn't fail does *not* mean you've got the answer right. There are hidden tests in your assignments. These check to ensure the code you've written does what we asked. As long as you do what we ask *and* pass the assert statements, you'll receive full credit.

<!-- BEGIN QUESTION -->

### Q2 - Math Question in $\LaTeX$

For each math question, you would need to use inline $\LaTeX$ in markdown to complete your answer. We understand not everyone is fluent in LaTeX but we think this is a good skill to learn. [Here is a fantastic tutorial from CalTech about using $\LaTeX$ in Jupyter Notebook.](http://chebe163.caltech.edu/2018w/handouts/intro_to_latex.html). You could also find various $\LaTeX$ tutorials and cheat sheets online. 

#### Example Question:
Given the following system of equations, what are the values of $x$ and $y$

$$
\begin{cases}
      x + y = 10\\
      x - y = 0
\end{cases}\,.
$$

_Sol:_ $x = 5, y = 5$

### Question
States Bayse Rule using events $A$ and $B$.

_Points:_ 0.1


$$ 
P(A \mid B) = \frac{P(B \mid A) \cdot P(A)}{P(B)}
$$

<!-- END QUESTION -->

### Q3 - Coding Question

You will be asked to either calculate a target value and assign the value to a specific variable, or to implement a function. Previously in Q1, you saw how to assign the target value to the variable. Here we shall see how to implement a function.

```python
import numpy as np

def euclidean_distance(pt1, pt2):
    """
    Given two numpy array, find the euclidean distance between
    the two points. Return a signle float.
    """
    difference = pt1 - pt2 # Element wise difference
    exponents = difference ** 2 # element wise exponent
    array_sum = np.sum(exponents)
    result = np.sqrt(array_sum)
    return result
```

```python
grader.check("q3_coding_question")
```

Output:
```
q3_coding_question results: All test cases passed!
```

<!-- BEGIN QUESTION -->

### Q4 - Free Response Question
 
Describe what you see. 

![image.png](imgs/racoon.jpeg)

_Points:_ 0.1

In my defense I was left unsupervised

<!-- END QUESTION -->

## The End!

This is the end of the practice assignment!

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. This restarts everything and runs your code from top to bottom.

When you are ready to submit your assignment, you can click `Validate` at the top. Note that in some assignments the code will take too long to run and validation may fail. Validation is just a final check that all the asserts are passing without failing.

Once you're happy with your work, click the disk icon to save, and submit the zip file onto gradescope. **You MUST submit all the required component to receive credit.**

Note that you can submit at any time, but **we grade your most recent submission**. This means that **if you submit an updated notebook after the submission deadline, it will be marked as late**.

_Note:_ If you encountered `LatexFailed` message during exporting, it indicates that your $\LaTeX$ code is not correct. Try to fix your LaTeX syntax error. If you need any additional helps, please make a private post on piazza and we are happy to help.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution.

```python
grader.export(force_save=True, run_tests=True, files=['imgs'])
```

Output:
```
<IPython.core.display.Javascript object>
Running your submission against local test cases...



Your submission received the following results when run against available test cases:

    q1_multiple_choice_question results: All test cases passed!

    q3_coding_question results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q1_multiple_choice_question (points: 0.1)
```
>>> possible_answers = ['A', 'B', 'C', 'D']
>>> assert favorite_class in possible_answers, 'Did you put in some other options?'
```


Tests for q3_coding_question (points: 0.1)
```
>>> assert callable(euclidean_distance)
```
```
>>> (pt1, pt2) = (np.array([1, 1]), np.array([0, 0]))
>>> assert np.isclose(np.sqrt(2), euclidean_distance(pt1, pt2))
```
